```python
# Initialize Otter
import otter
grader = otter.Notebook("mod01_Sicily.ipynb")
```

# Module 1: Intro to Human Geography!

### Welcome to `GEOG 102`

In this interactive notebook, we will work as geographic data scientists. A data scientist looks at data, asks questions, synthesizes and analyzes hypothesis', and makes predictions. The first step in this quest is to look at collected data, visualize the data, and observe. This will (hopefully) bring up all kinds of questions.

The goal of this activity is to review some of the course material from module 1 by:
- Importing data collected by the World Trade Organization (WTO)
- Focusing in on data of interest by looking at it in a tabular format
- Visualizing the data using graphs and maps
- Observing, ask questions, and learn from the data

We will ...
- Introduce jupyter notebooks as an interactive learning platform
- Review some concepts and terminology around maps
- Compare 2 countries (Italy and Congo), their location, history & uneven economic regional growth



# 1. What is a Jupyter notebook?
This webpage is called a Jupyter notebook. A notebook is a place to write code and view their results, and also to write text.
In a notebook, each rectangle containing text or code is called a *cell*.

**Text cells** (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings.
After you edit a text cell, click the "run cell" button at the top that looks like ▶| or hold down `shift` + `return` to confirm any changes. 

**Code cells** contain code in the Python 3 language. Running a code cell will execute all of the code it contains.
To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press ▶| or hold down `shift` + `return`.

**Activity 1:** This is a text cell. It is the cell type where we can type text that isn't code. Go ahead and double click in this cell and you will see that you can edit it. 

**Type something here:** ....


**Activity 2:** Click on the code cell below and run the code:

```python
#This cell is a code cell. It is where we can type code that can be executed.
#The hashtag at the start of this line makes it so that this text is a comment not code. 

print("Hello, World! \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

And this one:

```python
# This coding cell imports some python libraries that we will be using throughout this notebook
# Don't worry about what they are, just run this cell before running any other cells below this one
from datascience import *
import numpy as np
import otter
grader = otter.Notebook("mod01_Sicily.ipynb")

%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

import folium
def createMap(locationInit, zoomInit, marker1Location, marker1Name, marker2Location, marker2Name):
    myMap = folium.Map(location=locationInit, tiles="OpenStreetMap", zoom_start= zoomInit)
    folium.Marker(marker1Location, popup = marker1Name).add_to(myMap)
    folium.Marker(location = marker2Location, popup = marker2Name).add_to(myMap)
    return myMap

from IPython.display import YouTubeVideo
```

# 2. Comparing 2 locations and their history
Before you continue, make sure to review the following 2 lectures:
- Regional economic growth lecture! 
- Uneven economic development lecture

# 2.1 Maps
In this section we will start with a map. We will look at 2 countries locations Italy (focusing on Sicily) and the Democratic Republic of Congo. We will compare the difference in their geographies in terms of latitude, longitude and climate, and discuss how this affected their economic development. 

Both locations have similarities and differences when locating them on the map:

**Question 2.1** Let's first have a look at their locations on the map! In the coding cell below:
- Set **sicily_lat** and **sicily_long** equal to the latitude and longitude of Palermo (the capital city of Sicily, Southern Italy))
- Set **congo_lat** and **congo_long** equal to the latitude and longitude of Kinshasa (the capital of D.R., Congo)

HINT: You can use [Google Maps](https://www.google.com/maps) and if you have a PC, right click on a location to get the coordinates (latitude and longitude), if you have a Mac, single click with 2 fingers.


```python
#replace the ... with the correct latitudes and longitudes
sicily_lat = ...
sicily_long = ...
congo_lat = ...
congo_long = ...
createMap([33,0], 2, [sicily_lat, sicily_long], 'Palermo, Sicily', [congo_lat, congo_long], 'Kinshasa, Democratic Republic, Congo')
```

Now you can check your work by running the grader check below. If it passes, great job, go on to the next section! 

If it fails, don't worry, you just need to go back and try again.

```python
grader.check("q21")
```

**Question 2.2** Look at the map above. Which of the following statements is correct? 

Palermo (Sicily) and Kinshasa (D.R., Congo) have...
1. similar latitudes and similar longitudes.
2. similar latitudes, but different longitudes.
3. different latitudes, but similar longitudes.
4. different latitudes and different longitudes.

Answer in the code cell below, by setting the name (variable) `correct` equal to 1, 2, 3 or 4.

```python
correct = ...
```

```python
grader.check("q22")
```

Run the code cell below to watch the video "What caused the Rwandan Genocide?" and answer question 2.3.

```python
# The original URL is: "What caused the Rwandan Genocide"
# https://youtu.be/MF7EbUGlaOU

YouTubeVideo("MF7EbUGlaOU")
```

**Question 2.3:** What strongly impacted the Rwandan Genocide?
1. The legacy of colonialism 
2. Disagreements on farming techniques
3. An arms race with the DR Congo
4. all of these are correct

Answer in the code cell below, by setting the name (variable) `cause` equal to 1, 2, 3 or 4.

```python
cause = ...
```

```python
grader.check("q23")
```

# 2.2 Uneven Regional Economic Growth - Italy versus D.R. Congo

Let's explore data and compare Italy and the Democratic Republic of Congo by looking at merchandise export over the years. Do we see the same growth? Are their any data that stands out? Why?

## 2.2.1 History's Deadliest King

Before we look at data. Let's run the code cell below to watch a TEDEd talk on Belgian King Leopold. This video will help with the next question 2.4.

```python
# The original URL is: "Deadliest King"
# https://youtu.be/MU4vua2kNQY

YouTubeVideo("MU4vua2kNQY")
```

**Question 2.4:** After convincing Americans and Europeans he would continue _________, King Leopold was granted the Congo as his personal property:
1. the slave trade
2. free trade
3. rubber harvesting
4. searching for the source of the Nile

Answer in the code cell below, by setting the name (variable) `result` equal to 1, 2, 3 or 4.

```python
result = ...
```

```python
grader.check("q24")
```

Now let's get started with some **data exploration**!
## 2.2.2 Merchandise Export over the years
We are ready to import and visualize data. What can we learn about economies of Italy and D.R.Congo, by looking at merchandise export over the years?
The dataset used in this section is from the World Trade Organization (WTO): https://stats.wto.org/.

**STEP 1:** Gather the data from WTO, and upload it in our notebook as a table. 

Just run the cell below to see tabular data for both locations.

```python
# Data comparing Italy versus Congo merchandise export between 1948 and 2024: Exports in million US $ 
exports_data = Table.read_table('WtoData_20250528_export_Italy_Congo.csv')
exports_data
```

Take a moment to look at the table above, maybe we can see that there is growth, but it is very difficult to compare the growth trends by just looking at a table. The next step will help!

**STEP 2**: Visualize the table data by plotting the numbers for both locations. 
The cell below includes some helper code that will make plotting this information easier. Just run the cell and go to the next one.

```python
# This is a helper function used in the cells below, just run it (do not change anything!)
def plotExportGrowth(countryName, myTable, rowStart, rowEnd):
      ''' This function plots the export in Merchandise Export for a specific country '''
      myTable.plot('Year', countryName)
      plt.title('Merchandise Export (millions of $)\n'+ countryName)
      export_start = exports_data.column(countryName).item(rowStart)
      export_end = exports_data.column(countryName).item(rowEnd)
      growth = round(export_end/export_start, 2)
      print(countryName,'increased from '+ str(export_start) + ' to ' + 
            str(export_end)+' million $, growth factor = ' + str(growth))
```

The next coding cell will use the helper code above to calculate a growth factor and plot the data for Italy and Congo. Run the coding cell below, observe and note the differences between the 2 countries. The plots should help you answer question 2.5.

```python
print('Merchandise Export, 1948 to 2024:')
# Call the helper function to view data on Italy
plotExportGrowth('Italy', exports_data, 0, 76)
# Call the helper function to view data on Congo
plotExportGrowth('Congo', exports_data, 0, 76)
```

**Question 2.5:**
What does the data show? Do Italy and D.R., Congo show differences in merchandise export growth over the years? Are there any decades/years in particular that stand out? Why could this be? 

Review what you learned in lecture. In a couple of sentences, share your thoughts regarding your data observations and possible reasons.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**CONGRATULATIONS!** You just finished your first jupyter notebook assignment! 

We hope you enjoyed a first glimpse at the power of data and computer programming, tools that can help us understand how geography affects humans around the world!

Be sure to...

- run all of the tests and verify that they all pass, 
- check your answers for the open response questions,
- choose **Download as PDF** from the **File** menu,
- submit the .pdf file on **canvas**.

Tests for q21 (points: None)
```
>>> 38 <= sicily_lat <= 39
True
```
```
>>> 13 <= sicily_long <= 14
True
```
```
>>> -5 <= congo_lat <= -4
True
```
```
>>> 15 <= congo_long <= 16
True
```


Tests for q22 (points: None)
```
>>> correct == 3
True
```


Tests for q23 (points: None)
```
>>> cause == 1
True
```


Tests for q24 (points: None)
```
>>> result == 3
True
```
